> ***Consignes***
> Ce TP est à compléter directement en utilisant JupyterLab, version locale ou version en ligne https://jupyter.org/try
>
> *Remarque* Les personnes ayant un compte Google peuvent aussi utiliser https://colab.research.google.com/
>
> *Rappel* Supprimer les lignes "raise NotImplementedError()"
>
>Votre travail devra ensuite être
>- pour les G1, envoyé au format ipynb
>- pour les autres, envoyé au format html (File -> Export as ... ou Download as ...)

# Le modèle proie-prédateur (Modèle de Lotka-Voltera)

En dynamique de population, il est fréquent que les populations de deux espèces animales interagissent. Dans les années 1920-1930, l'italien Vito voltera et l'américain Alfred James Lotka ont mené séparément des recherches sur des systèmes où des prédateurs (requins ou lynx) et leurs proies (sardines en mer Adriatique ou rongueurs) sont liés, dans un milieu stable. 

Voici une version simplifiée de leur modèle.


Pour tout entier naturel $n$, on note $R_n$ l'effectif de la population de proies et $L_n$ celui de prédateurs au bout de $n$ années après l'introduction des prédateurs.

On fait les hypothèses que d'une année à l'autre

- le taux d'accroissement des proies est une fonction affine décroissante du nombre de prédateurs $$\forall n\in \mathbb{N}\,, \frac{R_{n+1}-R_n}{R_n} = a - b\, L_n \text{ avec } a\geqslant 0\,, b\geqslant 0 $$
- le taux d'accroissement des prédateurs est une fonction affine croissante du nombre de proies $$\forall n\in \mathbb{N}\,, \frac{L_{n+1}-L_n}{L_n} = c\,R_n - d \text{ avec } c\geqslant 0\,, d\geqslant 0 $$


## Application avec la population des Lynx et des petits rongueurs sur un territoire clos du Canada.

On suppose que le système donné ci-dessus se résume au cas particulier suivant $\begin{cases}L_{n+1}=0,96L_n+0,00005\times L_n\times R_n \\ R_{n+1} = 1,05\times R_n - 0,002\times L_n \times R_n \end{cases}$ 

On suppose aussi que $L_0 = 40$ et $R_0=2000$.

**Questions** En l'absence de prédateurs
1. Exprimer $R_{n+1}$ en fonction de $R_n$. Justifier.
2. Quelle est la nature de la suite $(R_n)$ ? Donner ses éléments caractéristiques.
3. Comment évolue l'effectif des proies ? Justifier.

<!-- grade_id: cell-59016b09924dc984, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

**Questions** En l'absence de rongeurs
1. Exprimer $L_{n+1}$ en fonction de $L_n$. Justifier.
2. Quelle est la nature de la suite $(L_n)$ ? Donner ses éléments caractéristiques.
3. Comment évolue l'effectif des prédateurs ? Justifier

<!-- grade_id: cell-92a17a3660cd7ab6, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-b2caa07e29c1b7af, points: 0, grade: false, solution: true -->
```python
'''Reproduire et compléter la fonction suivante prenant en argument un entier naturel n
et renvoyant les n premiers termes des suites (Ln) et (Rn)'''

'''
def lv(n) :
    l,r = 40,2000
    liste_lynx = []
    liste_rongeurs = ...
    for i in range(n):
        liste_lynx.append(l)
        liste_rongeurs.append(...)
        l,r = .96*l + .00005*l*r, ...
    return liste_lynx,liste_rongeurs
'''

def lv(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bb53095f7dc3aee4, points: 1, grade: true, solution: false -->
```python
''' Vérifier que votre fonction renvoie les valeurs pour différentes valeurs de n
    Cette cellule doit s'exécuter sans erreur '''

assert lv(2) == ([40, 42.4], [2000, 1940.0])
assert lv(3) == ([40, 42.4, 44.8168], [2000, 1940.0, 1872.488])
```

<!-- grade_id: cell-431f91041e5ae8ed, points: 0, grade: false, solution: true -->
```python
'''Reproduire et compléter la fonction suivante prenant en argument un entier naturel n
et représentant sur un même graphique les points (n;Ln) et (n;Rn)'''

%matplotlib inline
import matplotlib.pyplot as plt

'''
def glv(n):
    liste_x = ...
    liste_lynx, liste_rongeurs = lv(n)
    plt.plot(liste_x,liste_lynx,'ro', label='Lynx')
    plt.plot(...,...,'bo', label='Rongeurs')
    plt.xlabel("$n$")
    plt.title("Evolution de la population des rongeurs et des lynx en fonction du temps $n$")
    plt.legend()
    plt.show()
'''

def glv(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4485224627dfd5fa, points: 0, grade: false, solution: true -->
```python
''' Exécuter cette fonction pour obtenir 150 points de chaque suite'''

# YOUR CODE HERE
raise NotImplementedError()
```

**Questions** 
1. Décrire les variations de ces suites
2. En donner une interprétation

<!-- grade_id: cell-79cf5c019cc55afe, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-bb374d14a4a538bd, points: 0, grade: false, solution: true -->
```python
''' Exécuter la fonction précédente pour obtenir 1500 points de chaque suite'''

# YOUR CODE HERE
raise NotImplementedError()
```

**Questions** 
1. Décrire les variations de ces suites
2. En donner une interprétation

<!-- grade_id: cell-cd5260617e7695c2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd90efccfd79eb60, points: 0, grade: false, solution: true -->
```python
'''Reproduire et compléter la fonction suivante prenant en argument un entier naturel n
et représentant sur un même graphique les points (Ln;Rn)'''

'''
def nlv(n):
    liste_lynx, liste_rongeurs = ...
    plt.plot(...,...,'b>')
    plt.xlabel("Evolution de la population des lynx")
    plt.ylabel("Evolution de la population des rongeurs")
    plt.title("Evolution de la population des lynx et des rongeurs")
    plt.plot(40,2000,'ro', label='$(L_0;R_0)$')
    plt.legend()
    plt.show()
'''

def nlv(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9eea7c5088678a68, points: 0, grade: false, solution: true -->
```python
''' Exécuter cette fonction pour obtenir un nuage de 150 points'''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c506127a86c6278, points: 0, grade: false, solution: true -->
```python
''' Exécuter la fonction précédente pour obtenir un nuage de 1500 points'''

# YOUR CODE HERE
raise NotImplementedError()
```

```python

```